# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


```python
import pandas as pd 

df = pd.read_csv('assets/NISPUF17.csv', index_col = 0)
df
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT    PROVWT_D     RDDWT_D  STRATUM  YEAR  \
1       128521     12852     2         NaN  235.916956     1031  2017   
2        10741      1074     2         NaN  957.353840     1068  2017   
3       220011     22001     2         NaN  189.611299     1050  2017   
4        86131      8613     1  675.430817  333.447418     1040  2017   
5       227141     22714     1  482.617748  278.768063     1008  2017   
...        ...       ...   ...         ...         ...      ...   ...   
28461    19141      1914     2         NaN   21.188088     2062  2017   
28462   121401     12140     2         NaN   26.503010     2004  2017   
28463   128831     12883     2         NaN   27.520652     2004  2017   
28464    82921      8292     2         NaN   22.390587     2062  2017   
28465   244621     24462     2         NaN   27.520652     2004  2017   

       AGECPOXR  HAD_CPOX  AGEGRP  ...  XVRCTY2  XVRCTY3  XVRCTY4  XVRCTY5  \
1           NaN         2       1  ...              
```

```python
col1 = ["i0","i1","i2","i3","i4","i5"]
col2 = [1,2,1,2,2,3]
df = pd.DataFrame()
df["col1"] = col1
df["col2"] = col2
print(df)
print("--"*20)
print(df["col2"].value_counts())
```

Output:
```
col1  col2
0   i0     1
1   i1     2
2   i2     1
3   i3     2
4   i4     2
5   i5     3
----------------------------------------
2    0.500000
1    0.333333
3    0.166667
Name: col2, dtype: float64
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    import pandas as pd 
    df = pd.read_csv('assets/NISPUF17.csv', index_col = 0)
    medu = df['EDUC1']
    tlen = len(medu)
    PoE = {"less than high school":len(df[df['EDUC1']==1])/tlen,
           "high school":len(df[df['EDUC1']==2])/tlen,
           "more than high school but not college":len(df[df['EDUC1']==3])/tlen,
           "college":len(df[df['EDUC1']==4])/tlen}
    return PoE
```

```python
print(proportion_of_education())
```

Output:
```
{'less than high school': 0.10202002459160373, 'high school': 0.172352011241876, 'more than high school but not college': 0.24588090637625154, 'college': 0.47974705779026877}
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
import pandas as pd 
df = pd.read_csv('assets/NISPUF17.csv', index_col = 0)

print(df['CBF_01'].value_counts())
print(df['P_NUMFLU'].value_counts())
```

Output:
```
1     24261
2      4085
77      110
99        9
Name: CBF_01, dtype: int64
3.0    4845
0.0    4139
2.0    3350
1.0    1679
4.0    1282
5.0      35
6.0       3
Name: P_NUMFLU, dtype: int64
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    import pandas as pd 
    df = pd.read_csv('assets/NISPUF17.csv', index_col = 0)
    print("org len : ",len(df))
    df = df[df['CBF_01']<=2]
    print("new len : ",len(df))
    
    bm = df[df['CBF_01']==1]
    nbm = df[df['CBF_01']==2]
    print(len(bm), len(nbm))
    print("mean bm : ", np.mean(bm['P_NUMFLU']))
    print("mean nbm : ", np.mean(nbm['P_NUMFLU']))
    
    return ((bm['P_NUMFLU'].mean()),(nbm['P_NUMFLU'].mean()))
    # YOUR CODE HERE
    #raise NotImplementedError()
```

```python
import numpy as np
print (average_influenza_doses())
```

Output:
```
org len :  28465
new len :  28346
24261 4085
mean bm :  1.8799187420058687
mean nbm :  1.5963945918878317
(1.8799187420058687, 1.5963945918878317)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
org len :  28465
new len :  28346
24261 4085
mean bm :  1.8799187420058687
mean nbm :  1.5963945918878317
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    import pandas as pd 
    df = pd.read_csv('assets/NISPUF17.csv', index_col = 0)
    male = df[df['SEX']==1]
    female = df[df['SEX']==2]
    malevac = male[male['P_NUMVRC']>=1]
    femalevac = female[female['P_NUMVRC']>=1]
    malevac_cp = malevac[malevac['HAD_CPOX']==1]
    malevac_nocp = malevac[malevac['HAD_CPOX']==2]
    femalevac_cp = femalevac[femalevac['HAD_CPOX']==1]
    femalevac_nocp = femalevac[femalevac['HAD_CPOX']==2]
    ve_by_sex = {"male" : len(malevac_cp)/len(malevac_nocp),
                 "female" : len(femalevac_cp)/len(femalevac_nocp) }
    return ve_by_sex
    # YOUR CODE HERE
    #raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

```python
main = pd.read_csv('assets/NISPUF17.csv', index_col = 0)
print(main['HAD_CPOX'].unique())
print(main['P_NUMVRC'].unique())
```

Output:
```
[ 2  1 77 99]
[nan  1.  0.  2.  3.]
```

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    import pandas as pd 
    main = pd.read_csv('assets/NISPUF17.csv', index_col = 0)
    
    cpox = main[main['HAD_CPOX']<=2]
    cpox  = cpox[cpox['P_NUMVRC']>=0]
    # df=pd.DataFrame({"had_chickenpox_column" : cpox, "num_chickenpox_vaccine_column" : cpoxvac

    # here is some stub code to actually run the correlation
  #  corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    #return corr

    # YOUR CODE HERE
    corr, pval = stats.pearsonr(cpox['HAD_CPOX'], cpox['P_NUMVRC'])
    return corr
    #raise NotImplementedError()

```

```python
corr_chickenpox()
```

Output:
```
0.07044873460147986
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python

```